compile([[lit(pos,eq(fn(0,[var(3),fn(0,[var(2),var(1)])]),fn(0,[fn(0,[var(3),var(2)]),var(1)])))],[lit(pos,eq(fn(0,[fn(0,[fn(1,[]),fn(2,[])]),fn(1,[])]),fn(1,[])))],[lit(neg,eq(fn(0,[fn(0,[fn(1,[]),var(0)]),fn(1,[])]),fn(1,[]))),lit(neg,eq(fn(0,[fn(0,[var(0),fn(1,[])]),var(0)]),var(0)))]], [line(0,[lit(pos,eq(fn(0,[var(0),fn(0,[var(1),var(2)])]),fn(0,[fn(0,[var(0),var(1)]),var(2)])))],asm),line(1,[lit(pos,eq(fn(0,[fn(0,[fn(1,[]),fn(2,[])]),fn(1,[])]),fn(1,[])))],asm),line(2,[lit(neg,eq(fn(0,[fn(0,[fn(1,[]),var(3)]),fn(1,[])]),fn(1,[]))),lit(neg,eq(fn(0,[fn(0,[var(3),fn(1,[])]),var(3)]),var(3)))],asm),line(5,[lit(pos,eq(fn(0,[fn(1,[]),fn(0,[fn(2,[]),fn(1,[])])]),fn(1,[])))],rep(0,1)),line(9,[lit(pos,eq(fn(0,[fn(1,[]),var(0)]),fn(0,[fn(1,[]),fn(0,[fn(0,[fn(2,[]),fn(1,[])]),var(0)])])))],rep(0,5)),line(10,[lit(pos,eq(fn(0,[fn(1,[]),var(0)]),fn(0,[fn(1,[]),fn(0,[fn(2,[]),fn(0,[fn(1,[]),var(0)])])])))],rep(9,0)),line(13,[lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(1,[])])]),fn(1,[]))),lit(neg,eq(fn(0,[fn(0,[var(0),fn(1,[])]),var(0)]),var(0)))],rep(2,0)),line(23,[lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(0,[var(1),fn(1,[])])])]),fn(1,[]))),lit(neg,eq(fn(0,[var(0),var(1)]),fn(0,[fn(0,[fn(0,[var(0),var(1)]),fn(1,[])]),fn(0,[var(0),var(1)])])))],rep(13,0)),line(28,[lit(neg,eq(fn(0,[var(0),var(1)]),fn(0,[fn(0,[var(0),fn(0,[var(1),fn(1,[])])]),fn(0,[var(0),var(1)])]))),lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(0,[var(1),fn(1,[])])])]),fn(1,[])))],rep(23,0)),line(42,[lit(neg,eq(fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])]),fn(0,[fn(0,[var(0),fn(1,[])]),fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])])]))),lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(0,[fn(0,[fn(1,[]),fn(2,[])]),fn(1,[])])])]),fn(1,[])))],rep(28,1)),line(53,[lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(0,[fn(1,[]),fn(0,[fn(2,[]),fn(1,[])])])])]),fn(1,[]))),lit(neg,eq(fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])]),fn(0,[fn(0,[var(0),fn(1,[])]),fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])])])))],rep(42,0)),line(54,[lit(neg,eq(fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])]),fn(0,[fn(0,[var(0),fn(1,[])]),fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])])]))),lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(1,[])])]),fn(1,[])))],rep(53,5)),line(90,[lit(neg,eq(fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])]),fn(0,[var(0),fn(0,[fn(1,[]),fn(0,[var(0),fn(0,[fn(1,[]),fn(2,[])])])])]))),lit(neg,eq(fn(0,[fn(1,[]),fn(0,[var(0),fn(1,[])])]),fn(1,[])))],rep(54,0)),line(159,[lit(neg,eq(fn(0,[fn(2,[]),fn(0,[fn(1,[]),fn(2,[])])]),fn(0,[fn(2,[]),fn(0,[fn(1,[]),fn(2,[])])]))),lit(neg,eq(fn(0,[fn(1,[]),fn(0,[fn(2,[]),fn(1,[])])]),fn(1,[])))],rep(90,10)),line(162,[lit(neg,eq(fn(0,[fn(1,[]),fn(0,[fn(2,[]),fn(1,[])])]),fn(1,[])))],trv(159)),line(166,[],rsl(162,5))], X).


1. f0(X0,f0(X1,X2)) = f0(f0(X0,X1),X2) [input]
2. f0(f0(f1,f2),f1) = f1 [input]
3. f0(f0(f1,X3),f1) != f1 | f0(f0(X3,f1),X3) != X3 [input]
6. f0(f1,f0(f2,f1)) = f1 [superposition 1,2]
10. f0(f1,X0) = f0(f1,f0(f0(f2,f1),X0)) [superposition 1,6]
11. f0(f1,X0) = f0(f1,f0(f2,f0(f1,X0))) [forward demodulation 10,1]
14. f0(f1,f0(X0,f1)) != f1 | f0(f0(X0,f1),X0) != X0 [superposition 3,1]
24. f0(f1,f0(X0,f0(X1,f1))) != f1 | f0(X0,X1) != f0(f0(f0(X0,X1),f1),f0(X0,X1)) [superposition 14,1]
29. f0(X0,X1) != f0(f0(X0,f0(X1,f1)),f0(X0,X1)) | f0(f1,f0(X0,f0(X1,f1))) != f1 [forward demodulation 24,1]
43. f0(X0,f0(f1,f2)) != f0(f0(X0,f1),f0(X0,f0(f1,f2))) | f0(f1,f0(X0,f0(f0(f1,f2),f1))) != f1 [superposition 29,2]
54. f0(f1,f0(X0,f0(f1,f0(f2,f1)))) != f1 | f0(X0,f0(f1,f2)) != f0(f0(X0,f1),f0(X0,f0(f1,f2))) [forward demodulation 43,1]
55. f0(X0,f0(f1,f2)) != f0(f0(X0,f1),f0(X0,f0(f1,f2))) | f0(f1,f0(X0,f1)) != f1 [forward demodulation 54,6]
91. f0(X0,f0(f1,f2)) != f0(X0,f0(f1,f0(X0,f0(f1,f2)))) | f0(f1,f0(X0,f1)) != f1 [superposition 55,1]
160. f0(f2,f0(f1,f2)) != f0(f2,f0(f1,f2)) | f0(f1,f0(f2,f1)) != f1 [superposition 91,11]
163. f0(f1,f0(f2,f1)) != f1 [trivial inequality removal 160]
167. $false [subsumption resolution 163,6]]

f0(f2,f0(f1,f2)) != f0(f2,f0(f1,f0(f2,f0(f1,f2)))) | f0(f1,f0(f2,f1)) != f1 [superposition 55,1]
                          f0(f1,f0(f2,f0(f1,X0)))
                          f0(f1,X0) 
f0(f2,f0(f1,f2)) != f0(f2,f0(f1,f2)) | f0(f1,f0(f2,f1)) != f1 [superposition 91,11]













14. f0(f1,f0(X0,f1)) != f1 | f0(f0(X0,f1),X0) != X0 [superposition 3,1]

14. f0(f1,f0(X0,f1)) != f1 | X0 != f0(f0(X0,f1),X0) != X0 [superposition 3,1]

1. f0(X0,f0(X1,X2)) = f0(f0(X0,X1),X2) [input]
24. f0(f1,f0(X0,f0(X1,f1))) != f1 | f0(X0,X1) != f0(f0(f0(X0,X1),f1),f0(X0,X1)) [superposition 14,1]



14. f0(f1,f0(f0(X0,X1),f1)) != f1 | f0(X0,X1) != f0(f0(f0(X0,X1),f1),f0(X0,X1)) [superposition 3,1]
          f0(f0(X0,X1),f1)                          f0(f0(X0,X1),f1) 
          f0(X0,f0(X1,f1)) 
24. f0(f1,f0(X0,f0(X1,f1))) != f1 | f0(X0,X1) != f0(f0(f0(X0,X1),f1),f0(X0,X1)) [superposition 14,1]